#############################################
###       Generated from PreTeXt source      
###       on 2024-08-01T17:56:24-07:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.1.6.1)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "scaffold.pl",
  "AnswerFormatHelp.pl",
  "parserOneOf.pl",
  "contextArbitraryString.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');
COMMENT('This problem is scaffolded with multiple parts');

############################################################
# Header
############################################################
TEXT(beginproblem());


############################################################
# PG Setup Code
############################################################
Context()->variables->are();
for my $i (1..26) {
  Context()->variables->add($ALPHABET[$i-1] => 'Real');
  Context()->variables->add(lc($ALPHABET[$i-1]) => 'Real');
};
$d = OneOf('d', 'D', 'x', 'y');
$w = OneOf('w', 'W', 'x', 'y');
Context("ArbitraryString");
$ftorm = Compute("feet")->cmp(
  mathQuillOpts => 'disabled',
  checker => sub {
    my ( $correct, $student, $ansHash ) = @_;
    return 1 if ($student =~ /^(feet|foot|ft|ft~~.|meters|meter|m|m~~.)$/);
    return 0;
  }
);
$lborkg = Compute("pounds")->cmp(
  mathQuillOpts => 'disabled',
  checker => sub {
    my ( $correct, $student, $ansHash ) = @_;
    return 1 if ($student =~ /^(pounds|pound|lb|lbs|lb~~.|lbs~~.|kilograms|kilogram|kg|kgs|kg~~.|kgs~~.)$/);
    return 0;
  }
);

############################################################
# Body
############################################################

BEGIN_PGML
Identify a variable you might use to represent each quantity. Then identify what units would be most appropriate.


END_PGML

############################################################
# Scaffold
############################################################
Scaffold::Begin(numbered => 1,is_open => "correct_or_first_incorrect",can_open => always,);

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Let [_]{$d}{3} be the depth of a swimming pool, measured in [_]{$ftorm}{5}.


END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
The unknown quantity is depth, which starts with "d". We generally measure depth in feet for a swimming pool. (Meters is another reasonable unit.) So we could define this variable as:

"Let [`d`] be the depth of a swimming pool, measured in feet."


END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Let [_]{$w}{3} be the weight of a dog, measured in [_]{$lborkg}{5}.


END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
The weight of the dog is the unknown quantity, and "weight" starts with "w". We generally measure the weight of a dog in pounds. (Kilograms is another reasonable unit.) So we could define this variable as:

"Let [`w`] be the weight of a dog, measured in pounds."


END_PGML_SOLUTION

Section::End();

Scaffold::End();

############################################################
# End Problem
############################################################

ENDDOCUMENT();
