#############################################
###       Generated from PreTeXt source      
###       on 2024-08-11T14:51:26-07:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.3.4.35)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PGchoicemacros.pl",
  "contextArbitraryString.pl",
  "contextFraction.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());


############################################################
# PG Setup Code
############################################################
Context("Fraction");
$d = list_random(3,6,7,9,11,12,13,14);
($d, $n) = random_coprime([$d], [1..2*$d]);
$a = Fraction($n, $d);
$b = Real(Round($n/$d, 2));
Context("ArbitraryString");
($a, $b) = ($a, $b)[NchooseK(2, 2)];
$operator = ($a > $b) ? '>' : '<';
$answer = Compute($operator)->cmp(
    checker => sub {
        my ($correct,$student,$ans) = @_;
        $correct = $correct->value;
        $student = $student->value;
        if ($student =~ /^~~s*$correct~~s*$/) {return 1}
        elsif ($student !~ /^~~s*>~~s*$/ && $student !~ /^~~s*<~~s*$/ && $student !~ /^~~s*=~~s*$/)
            {Value->Error("Answer with >, <, or =.")}
        else {return 0};
    }
);

############################################################
# Body
############################################################

BEGIN_PGML
Decide if one given number is greater than, less than, or equal to another given number.

[`[$a]`] [_]{$answer}{3} [`[$b]`]


END_PGML

############################################################
# End Problem
############################################################

ENDDOCUMENT();
