#############################################
###       Generated from PreTeXt source      
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.7.6.23)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

#############
# Load Macros
#############

loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "contextForm.pl",
  "ORCCA.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

#######
# Setup
#######

@var = random_subset(3, 'a'..'n', 'p'..'z', 'A'..'D', 'F'..'N', 'P'..'T', 'V'..'Z');
Context("Numeric")->variables->are($var[0] => 'Real', $var[1] => 'Real', $var[2] => 'Real');
$given = Formula("$var[0] $var[1] + $var[2]");
$ans = Formula("$var[1] $var[0] + $var[2]");
$evaluator = $ans->cmp(checker => sub{
  my ($correct, $student, $ans) = @_;
  Value->Error("This is the original expression") if ($student eq $given);
  Value->Error("This is algebraically equivalent to the given expression, but using commutativity of multiplication would give you something else") if (($student == $correct) && ($student ne $correct));
  return ($student eq $correct);
});

#############################
# Exercise Group Introduction
#############################

BEGIN_PGML
Demonstrate that you know the meanings of the various algebraic properties by applying the given algebraic property to the given expression to get a new expression.

END_PGML

######
# Body
######

BEGIN_PGML
Apply commutativity of multiplication to [`[$given]\text{.}`]

[_]{$evaluator}{16}

END_PGML

##########
# Solution
##########

BEGIN_PGML_SOLUTION
There is one place in this expression where there is multiplication. To apply commutativity of multiplication, we just write those two factors in the opposite order. This gives us [`[$ans]\text{.}`]

END_PGML_SOLUTION

#############
# End Problem
#############

ENDDOCUMENT();
