#############################################
###       Generated from PreTeXt source      
###       on 2024-08-11T14:51:27-07:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(1.8.6.2)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "scaffold.pl",
  "AnswerFormatHelp.pl",
  "parserOneOf.pl",
  "contextArbitraryString.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');
COMMENT('This problem is scaffolded with multiple parts');

############################################################
# Header
############################################################
TEXT(beginproblem());


############################################################
# PG Setup Code
############################################################
Context()->variables->are();
for my $i (1..26) {
  Context()->variables->add($ALPHABET[$i-1] => 'Real');
  Context()->variables->add(lc($ALPHABET[$i-1]) => 'Real');
};
$a=OneOf('a','A','x','y','z');
$d=OneOf('d','D','m','M','x','y','z');
Context("ArbitraryString");
$years = Compute("years")->cmp(
  mathQuillOpts => 'disabled',
  checker => sub {
    my ( $correct, $student, $ansHash ) = @_;
    return 1 if ($student =~ /^(years|yrs|yr|y|yrs~~.|yr~~.|y~~.)$/);
    return 0;
  }
);
$miles = Compute("miles")->cmp(
  mathQuillOpts => 'disabled',
  checker => sub {
    my ( $correct, $student, $ansHash ) = @_;
    return 1 if ($student =~ /^(miles|mile|mi|mis|mis~~.|mi~~.)$/);
    return 0;
  }
);
$sqftorsqm = Compute("ft^2 or m^2")->cmp(
  mathQuillOpts => 'disabled',
  checker => sub {
    my ( $correct, $student, $ansHash ) = @_;
    return 1 if ($student =~ /^(ft~~^2|square feet|feet squared|feet square|sq ft|sqft|sq~~. ft~~.|ft~~.~~^2|m~~^2|square meters|meters squared|meters square|sq m|sqm|sq~~. m~~.|m~~.~~^2)$/);
    return 0;
  }
);

############################################################
# Body
############################################################

BEGIN_PGML
Identify a variable you might use to represent each quantity. Then identify what units would be most appropriate.


END_PGML

############################################################
# Scaffold
############################################################
Scaffold::Begin(numbered => 1,is_open => "correct_or_first_incorrect",can_open => always,);

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Let [_]{$a}{2} be the age of a person, measured in [_]{'years'}{10}.


END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
The unknown quantity is age, which we generally measure in years. So we could define this variable as:

"Let [`a`] be the age of a person, measured in years."


END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Let [_]{$d}{2} be the distance traveled by a driver that left Portland, Oregon, bound for Boise, Idaho, measured in [_]{'miles'}{10}.


END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
The distance traveled is the unknown quantity. Since this is a drive from Portland to Boise, it would make sense to measure this in miles. So we could define this variable as:

"Let [`d`] be the distance traveled by a driver that left Portland, Oregon, bound for Boise, Idaho, measured in miles."


END_PGML_SOLUTION

Section::End();

############################################################
# Section
############################################################
Section::Begin("");

BEGIN_PGML
Let [_]{$a}{2} be the surface area of the walls of a room, measured in [_]{'ft^2'}{10}.


END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
The unknown quantity is area. Apartment area is usually measured in square feet. So we'll define this variable as:

"Let [`A`] be the area of a two-bedroom apartment, measured in [`\text{ft}^2\text{.}`]"


END_PGML_SOLUTION

Section::End();

Scaffold::End();

############################################################
# End Problem
############################################################

ENDDOCUMENT();
