#############################################
###       Generated from PreTeXt source      
###       on 2024-08-01T17:56:28-07:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(3.5.18)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "bizarroArithmetic.pl",
  "AnswerFormatHelp.pl",
  "parserAssignment.pl",
  "contextForm.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());


############################################################
# PG Setup Code
############################################################
Context()->operators->set(
  '+'=>{class=>'bizarro::BOP::add',isCommand=>1},
  '-'=>{class=>'bizarro::BOP::subtract',isCommand=>1},
);
parser::Assignment->Allow;
Context()->variables->add(y=>'Real');
$eq=Formula("y=-40x+30");
$ev=$eq->cmp(checker=>sub{
my($co,$st,$aH)=@_;
return 0 if$aH->{isPreview}||$co!=$st;
$st=$aH->{student_formula};
$co=$co->{original_formula} if defined $co->{original_formula};
$st=Formula("$st");$co=Formula("$co");
return 0 unless($co==$st);
Context()->flags->set(bizarroAdd=>1,bizarroSub=>1);
delete $co->{test_values},$st->{test_values};
Value->Error("Your answer is an equation for this line, but it is not in simplified slope-intercept form")unless(($co==$st)or($st==$co));
Context()->flags->set(bizarroAdd=>0,bizarroSub=>0);
return 1;
})->withPostFilter(sub{my$ans=shift;
$ans->{ans_message}="An equation for a line in slope-intercept form should start with 'y='"
if $ans->{ans_message} eq "The left side of an assignment must be a variable or function";
return $ans;
});

############################################################
# Body
############################################################

BEGIN_PGML
Find the slope-intercept form of an equation for the line that passes through the points [`(-3,150)`] and [`(0,30)\text{.}`]

[_]{$ev}{10}


END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
We first find the slope between our two points: [`(\overset{x_1}{-3},\overset{y_1}{150})`] and [`(\overset{x_2}{0},\overset{y_2}{30})\text{.}`] Using the slope formula, we have:

[```\newcommand{\substitute}[1]{{\color{blue}{{#1}}}}\newcommand{\amp}{&}
\begin{aligned}
m\amp=\frac{y_2-y_1}{x_2-x_1}\\
\amp=\frac{\substitute{30}-\substitute{150}}{\substitute{0}-\substitute{(-3)}}\\
\amp=\frac{-120}{3}\\
\amp=-40
\end{aligned}
```]



Now we can write [`y=-40x+b`] and to find [`b`] we need look no further than one of the given points: [`(0,30)\text{.}`] The value of [`b`] is [`30\text{.}`] So, the slope-intercept form of the line is [`y=-40x+30\text{.}`]


END_PGML_SOLUTION

############################################################
# End Problem
############################################################

ENDDOCUMENT();
