#############################################
###       Generated from PreTeXt source      
###       on 2024-08-11T14:51:31-07:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(4.1.22)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PGlateximage.pl",
  "PCCmacros.pl",
  "parserGraphTool.pl",
  "ORCCA.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());


############################################################
# PG Setup Code
############################################################
$gt = GraphTool(
  "{line, solid, (0, -5), (8, 0)}",
)->with(
  availableTools => [ "LineTool" ],
  showInStatic => 0,
  cmpOptions => {
    list_checker => sub {
      my ($correct, $student, $ans, $value) = @_;
      return 0 if $ans->{isPreview};
      my $score = 0; my @errors; my $count = 1;
      my ($p1x, $p1y) = $correct->[0]->extract(3)->value;
      my ($p2x, $p2y) = $correct->[0]->extract(4)->value;
      my @stdform = (
        $p1y - $p2y, $p2x - $p1x, $p1x * $p2y - $p2x * $p1y,
      );
      my $pointOnLine = sub {
        my ($a, $b, $c, $point) = @_;  my ($x, $y) = $point->value;
        return $a * $x + $b * $y + $c == 0;
      };
      my @line = (0);
      for (@$student) {
        my $nth = Value::List->NameForNumber($count++);
        if ($pointOnLine->(@stdform[0..2],$_->extract(3)) && $pointOnLine->(@stdform[0..2],$_->extract(4))) {
          $line[0] = 1; $score += 1; next
        }
        push(@errors, "The $nth line graphed is incorrect.");
      }
      return (($line[0]) ? $score : 0, @errors);
    }
  }
);
Context("Point");
Context()->strings->remove("NONE");
Context()->strings->add(
  "no solutions" => {},
);
for (
  "no solution",
  "none",
  "there are no solutions",
  "the solution set is empty",
  "solution set is empty",
  "empty set",
  "∅",
  "{}",
  "{ }",
  "{  }",
  "{   }"
) {
  Context()->strings->add($_ => {alias => "no solutions"});
}
Context()->strings->add(
  "infinitely many solutions" => {},
);
for (
  "infinitely many",
  "infinite solution set",
  "infinite",
  "the solution set is infinite",
  "solution set is infinite"
) {
  Context()->strings->add($_ => {alias => "infinitely many solutions"});
}
$solution = Compute("infinitely many solutions")->cmp(
  typeMatch => Point(1,2),
  mathQuillOpts => 'disabled'
);
$image_1 = createLaTeXImage();
$image_1->addToPreamble(latexImagePreamble());
$image_1->BEGIN_LATEX_IMAGE
\begin{tikzpicture}
  \begin{axis}[
      width=0.47\orccaprintwidth
    ]
    \addplot[firstcurve, domain=-3:7]{5/8*x-5.05} node[pos=1,sloped,above left] {\(5x-8y=40\)};
    \addplot[secondcurve, domain=-3:7]{5/8*x-4.95} node[pos=0.9,sloped,below left] {\(y=\frac{5}{8}(x+8)-10\)};
  \end{axis}
\end{tikzpicture}

END_LATEX_IMAGE

############################################################
# Body
############################################################

BEGIN_PGML
Solve the following system of equations by graphing.

[```\newcommand{\amp}{&}
              \left\{
                \begin{alignedat}{3}
                  \amp5x-8y = 40\\
                  \amp y=\tfrac{5}{8}(x+8)-10
                \end{alignedat}
              \right.
            ```]



[_]{$gt}{1}

[@KeyboardInstructions('Enter the solution here. If there are no solutions, enter [|no solutions|]*. If there are infinitely many solutions, enter [|infinitely many solutions|]*.')@]**

[_]{$solution}{16}


END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
If we graph these lines, we find they are the same line.

>> [@image(insertGraph($image_1), width=>282, alt=>qq"a coordinate plot of two lines that are coinciding; they each have a y-intercept of (0,-5) and an x-intercept of (8,0)")@]*  <<

So there are infinitely many solutions: all points [`(x_1,y_1)`] on that common line.


END_PGML_SOLUTION

############################################################
# End Problem
############################################################

ENDDOCUMENT();
