#############################################
###       Generated from PreTeXt source      
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(4.1.5.44)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

#############
# Load Macros
#############

loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PGlateximage.pl",
  "scaffold.pl",
  "PCCmacros.pl",
  "parserMultiAnswer.pl",
  "parserNumberWithUnits.pl",
  "contextForm.pl",
  "parserLinearRelation.pl",
  "parserGraphTool.pl",
  "ORCCA.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');
COMMENT('This problem is scaffolded with multiple parts');

#######
# Setup
#######

Context("LinearRelation")->variables->are(t=>'Real',d=>'Real');
$rate[1]=random(30,45,5);
$dogspeed=NumberWithUnits("$rate[1] ft/s");
$x=list_random(4,8,10);
$time=NumberWithUnits("$x s");
$y=$x*$rate[1];
$distance=NumberWithUnits("$y ft");
$b=random(30,50,10);
$initial=NumberWithUnits("$b ft");
$rate[0]=($y-$b)/$x;
$speed=NumberWithUnits("$rate[0] ft/s");
@eq=(LinearRelation("d=$rate[0]t+$b"),LinearRelation("d=$rate[1]t"));
Context("Numeric");
@f = (Formula("$rate[0]*x+$b"), Formula("$rate[1]*x"));
$m=MultiAnswer(@eq)->with(singleResult=>0,
checker=>sub{
my($c,$s,$self)=@_;
my @s=@{$s};
my @c=@{$c};
if (($c[0]==$s[0]&&$c[1]==$s[1])||($c[0]==$s[1]&&$c[1]==$s[0])){return [1,1]}
elsif ($c[0]==$s[0]||$c[1]==$s[0]){return [1,0]}
elsif ($c[0]==$s[1]||$c[1]==$s[1]){return [0,1]}
else {return [0,0]}
});
$ymax=int($y/100+0.999)*100;
$gt=GraphTool("{line,solid,(0,$b),(1,$b+$rate[0])}","{line,solid,(0,0),(1,$rate[1])}")->with(
availableTools=>["LineTool"],showInStatic=>0,bBox=>[-2,$ymax,12,-20],snapSizeY=>10,gridY=>10,
ticksDistanceY=>50,minorTicksY=>4,xAxisLabel=>'t',yAxisLabel=>'d',
cmpOptions=>{list_checker=>sub{
my($c,$s,$a,$v)=@_;
return 0 if $ans->{isPreview};
my $score=0; my @errors; my $count=1;
my ($x1,$y1)=$c->[0]->extract(3)->value;
my ($x2,$y2)=$c->[0]->extract(4)->value;
my ($x3,$y3)=$c->[1]->extract(3)->value;
my ($x4,$y4)=$c->[1]->extract(4)->value;
my @std=($y1-$y2,$x2-$x1,$x1*$y2-$x2*$y1,$y3-$y4,$x4-$x3,$x3*$y4-$x4*$y3);
my $p=sub{my ($A,$B,$C,$pt)=@_; my ($x,$y)=$pt->value; return $A*$x+$B*$y+$C==0;};
my @l=(0,0);
for (@$s) {
my $nth=Value::List->NameForNumber($count++);
if ($p->(@std[0..2],$_->extract(3))&&$p->(@std[0..2],$_->extract(4))) {$l[0]=1;$score += 1; next}
if ($p->(@std[3..5],$_->extract(3))&&$p->(@std[3..5],$_->extract(4))) {$l[1]=1;$score += 1; next}
push(@errors, "The $nth line graphed is incorrect.");
}
return (($l[0]&&$l[1]) ? $score : 0,@errors);
}});

$image_1 = createLaTeXImage();
$image_1->addToPreamble(latexImagePreamble());
$image_1->BEGIN_LATEX_IMAGE
\begin{tikzpicture}
\begin{axis}[xmin=0,xmax=12,ymin=0,ymax=300,xtick={0,4,8,12},ytick={0,50,...,$ymax},minor xtick={0,1,...,12},minor ytick={0,10,...,$ymax},xlabel={\(t\)},ylabel={\(d\)}]
\addplot[firstcurve,-{Kite},domain=0:12]{$f[0]};
\addplot[secondcurve,-{Kite},domain=0:12]{$f[1]};
\end{axis}
\end{tikzpicture}

END_LATEX_IMAGE

######
# Body
######

BEGIN_PGML
A cyclist riding at [`[$speed]`] rides past a dog. A moment later, when the bicycle is [`[$initial]`] away, the dog begins to chase the bicycle at a speed of [`[$dogspeed]\text{.}`] Assuming the cyclist does not change course or speed, how long will it be until the dog catches up to the bicycle?

END_PGML

##########
# Scaffold
##########

Scaffold::Begin(numbered => 1,is_open => "correct_or_first_incorrect",can_open => always,);

#########
# Section
#########

Section::Begin("");

BEGIN_PGML
Write two equations that form a system for this scenario.

[@KeyboardInstructions('Use [|t|]* as the variable for how much time has passed, and [|d|]* for the variable representing how many feet the bicycle/dog has traveled.')@]**

[_]{$m}{16}

[_]{$m}{16}

END_PGML

##########
# Solution
##########

BEGIN_PGML_SOLUTION
The bicycle starts out [`[$initial]`] away and is traveling at [`[$speed]\text{.}`] So one equation is [`[$eq[0]]\text{.}`]

The dog starts running at time [`t=0`] with speed [`[$dogspeed]\text{.}`] So the other equation is [`[$eq[1]]\text{.}`]

END_PGML_SOLUTION

Section::End();

#########
# Section
#########

Section::Begin("");

BEGIN_PGML
Plot the lines for the system of two equations.

[_]{$gt}{1}

END_PGML

##########
# Solution
##########

BEGIN_PGML_SOLUTION
Plotting the two slope-intercept equations:

>> [@image(insertGraph($image_1), width=>282, alt=>qq"a Cartesian grid with two intersecting lines; one line passes through (0,$b) with slope $rate[0]; the other line passes through (0,0) with slope $rate[1]; the lines cross at ($x,$y)")@]*  <<

END_PGML_SOLUTION

Section::End();

#########
# Section
#########

Section::Begin("");

BEGIN_PGML
Based on the graph, how long will it be until the dog catches up with the bicycle?

[_]{$time}{16}

How far will that be from the place where the dog started running?

[_]{$distance}{16}

END_PGML

##########
# Solution
##########

BEGIN_PGML_SOLUTION
The lines cross at [`([$x],[$y])\text{.}`] So after [`[$x]`] seconds, the cyclist and the dog are each [$y] feet away from where the dog started running.

END_PGML_SOLUTION

Section::End();

Scaffold::End();

#############
# End Problem
#############

ENDDOCUMENT();
