#############################################
###       Generated from PreTeXt source      
###       on 2024-08-01T17:56:29-07:00       
###   A recent stable commit (2022-07-01):   
### 6c761d3dba23af92cba35001c852aac04ae99a5f 
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(4.1.12)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

############################################################
# Load Macros
############################################################
loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PGlateximage.pl",
  "PCCmacros.pl",
  "parserGraphTool.pl",
  "ORCCA.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

############################################################
# Header
############################################################
TEXT(beginproblem());


############################################################
# PG Setup Code
############################################################
$gt = GraphTool(
  "{line, solid, (0, 4), (-12, 0)}",
  "{line, solid, (0, -4), (-4, 0)}"
)->with(
  availableTools => [ "LineTool" ],
  showInStatic => 0,
  bBox => [-15, 15, 15, -15],
  ticksDistanceX => 5,
  ticksDistanceY => 5,
  minorTicksX => 0,
  minorTicksY => 0,
  cmpOptions => {
    list_checker => sub {
      my ($correct, $student, $ans, $value) = @_;
      return 0 if $ans->{isPreview};
      my $score = 0; my @errors; my $count = 1;
      my ($p1x, $p1y) = $correct->[0]->extract(3)->value;
      my ($p2x, $p2y) = $correct->[0]->extract(4)->value;
      my ($p3x, $p3y) = $correct->[1]->extract(3)->value;
      my ($p4x, $p4y) = $correct->[1]->extract(4)->value;
      my @stdform = (
        $p1y - $p2y, $p2x - $p1x, $p1x * $p2y - $p2x * $p1y,
        $p3y - $p4y, $p4x - $p3x, $p3x * $p4y - $p4x * $p3y
      );
      my $pointOnLine = sub {
        my ($a, $b, $c, $point) = @_;  my ($x, $y) = $point->value;
        return $a * $x + $b * $y + $c == 0;
      };
      my @line = (0, 0);
      for (@$student) {
        my $nth = Value::List->NameForNumber($count++);
        if ($pointOnLine->(@stdform[0..2],$_->extract(3)) && $pointOnLine->(@stdform[0..2],$_->extract(4))) {
          $line[0] = 1; $score += 1; next
        }
        if ($pointOnLine->(@stdform[3..5],$_->extract(3)) && $pointOnLine->(@stdform[3..5],$_->extract(4))) {
          $line[1] = 1; $score += 1; next
        }
        push(@errors, "The $nth line graphed is incorrect.");
      }
      return (($line[0] && $line[1]) ? $score : 0, @errors);
    }
  }
);
Context("Point");
$solution = Point(-6,2);
$image_1 = createLaTeXImage();
$image_1->addToPreamble(latexImagePreamble());
$image_1->BEGIN_LATEX_IMAGE
\begin{tikzpicture}
  \begin{axis}[
      xmin = -15,
      xmax = 15,
      ymin = -15,
      ymax = 15,
      xtick = {-15,-10,...,15},
      ytick = {-15,-10,...,15},
      minor xtick = {-15,-14,...,15},
      minor ytick = {-15,-14,...,15},
      width = 0.47\orccaprintwidth
    ]
    \addplot[firstcurve, domain=-15:15]{1/3*x+4} node[pos=1,sloped,above left] {\(x-3y=-12\)};
    \addplot[secondcurve, domain=-15:11]{-x-4} node[pos=1,sloped,above left] {\(x+y=-4\)};
    \addplot[soliddot,color=firstcolor] coordinates {(-12,0) (0, 4)};
    \addplot[soliddot,color=secondcolor] coordinates {(-4,0) (0,-4)};
  \end{axis}
\end{tikzpicture}

END_LATEX_IMAGE

############################################################
# Body
############################################################

BEGIN_PGML
Solve the following system of equations by graphing. Note that the equations are in standard form. To review plotting a line equation that is in standard form, see Section[$NBSP]*3.7.

[```\newcommand{\amp}{&}
              \left\{
                \begin{alignedat}{4}
                  x \amp {}-{} \amp 3y \amp {}={} \amp {-12} \\
                  x \amp {}+{} \amp y \amp {}={} \amp {-4}
                \end{alignedat}
              \right.
            ```]



[_]{$gt}{1}

[@KeyboardInstructions('Enter the solution here.')@]**

[_]{$solution}{5}


END_PGML

############################################################
# Solution
############################################################

BEGIN_PGML_SOLUTION
Both of these equations are written in standard form. The first equation, [`x-3y=-12\text{,}`] has [`x`]-intercept at [`(-12,0)`] and [`y`]-intercept at [`(0,4)\text{.}`] The second equation, [`x+y=-4\text{,}`] has [`x`]-intercept at [`(-4,0)`] and [`y`]-intercept at [`(0,-4)\text{.}`] We can use this information to graph both lines.

>> [@image(insertGraph($image_1), width=>282, alt=>qq"a Cartesian grid with two intersecting lines; one line has a x-intercept of -12 and y-intercept of 4; the other line has x-intercept of -4 and y-intercept of -4")@]*  <<

It appears that the two lines intersect where [`x=-6`] and [`y=2\text{,}`] so the solution to the system of equations would be the point [`(-6,2)\text{.}`] However we should check that [`(-6,2)`] actually works as a solution to each of the original equations.

[```\newcommand{\substitute}[1]{{\color{blue}{{#1}}}}\newcommand{\confirm}[1]{\stackrel{\checkmark}{#1}}\newcommand{\wonder}[1]{\stackrel{\text{?}}{#1}}\newcommand{\amp}{&}
\begin{aligned}
x-3y\amp=-12\amp x+y\amp=-4\\
\substitute{-6}-3(\substitute{2})\amp\wonder{=}-12\amp \substitute{-6}+\substitute{2}\amp\wonder{=}-4\\
-6-6\amp\confirm{=}-12\amp -4\amp\confirm{=}-4
\end{aligned}
```]

This verifies that [`(-6,2)`] is the solution.


END_PGML_SOLUTION

############################################################
# End Problem
############################################################

ENDDOCUMENT();
