#############################################
###       Generated from PreTeXt source      
###                                          
###         https://pretextbook.org          
###                                          
#############################################
## DBsubject()
## DBchapter()
## DBsection()
## Level()
## KEYWORDS()
## TitleText1()
## EditionText1()
## AuthorText1()
## Section1(not reported)
## Problem1(4.2.5.62)
## Author()
## Institution()
## Language(en-US)

DOCUMENT();

#############
# Load Macros
#############

loadMacros(
  "PGstandard.pl",
  "PGML.pl",
  "AnswerFormatHelp.pl",
  "PCCmacros.pl",
  "parserMultiAnswer.pl",
  "parserNumberWithUnits.pl",
  "contextCurrency.pl",
  "parserLinearRelation.pl",
  "ORCCA.pl",
  "PGcourse.pl",
);
COMMENT('Authored in PreTeXt');

#######
# Setup
#######

$r1 = random(1.60,1.80,0.01);
$r2 = random(5.60,6.00,0.01);
$r3 = random(2.50,3,0.10);
Context("Currency")->flags->set(trimTrailingZeros => 1);
$rate1 = Currency($r1);
$rate2 = Currency($r2);
$rate3 = Currency($r3);
Context("LinearRelation")->variables->are(P => 'Real', C => 'Real');
@eq = (LinearRelation("P + C = 1"), LinearRelation("$r1 P + $r2 C = $r3"));
$m = MultiAnswer(@eq)->with(
  singleResult => 0,
  checker => sub {
      my ( $c, $s, $self ) = @_;
      return [1,1] if ($c->[0] == $s->[0] && $c->[1] == $s->[1] || $c->[1] == $s->[0] && $c->[0] == $s->[1]);
      return [1,0] if ($c->[0] == $s->[0] || $c->[1] == $s->[0]);
      return [0,1] if ($c->[0] == $s->[1] || $c->[1] == $s->[1]);
      return [0,0];
  }
);
$c = ($r1 - $r3)/($r1 - $r2);
$p = 1 - $c;
Context("Numeric")->flags->set(tolerance => 2);
$C = NumberWithUnits("$c lb");
$P = NumberWithUnits("$p lb");

######
# Body
######

BEGIN_PGML
A snack company will produce and sell 1-lb bags with a mixture of peanuts and cashews. Peanuts cost [`[$rate1]`] per pound, and cashews cost [`[$rate2]`] per pound. The company is targeting a product that will cost them [`[$rate3]`] per pound worth of ingredients. How much of each type of nut should go into a bag?

[@KeyboardInstructions('Using [`P`] and [`C`] as the variables, give two equations that form a system of equations that is appropriate for this exercise.')@]**

[_]{$m}{16}

[_]{$m}{16}

[@KeyboardInstructions('Solve the system and enter the amount of peanuts here.')@]**

[_]{$P}{16}

[@KeyboardInstructions('Enter the amount of cashews here.')@]**

[_]{$C}{16}

END_PGML

#############
# End Problem
#############

ENDDOCUMENT();
